# Lista 00 - Revisão Python e Numpy

## Principalmente útil para quem não fez a versão nova de ALC!!!

# Exercício 01:

Crie uma função que, dado dois textos, retorne uma __lista__ (que pode ser vazia) em ordem lexicográfica com as palavras que ocorrem em ambos. A função não deve ser case sensitive, nesse caso, transforme todo o texto reescrevendo as palavras em letra minúscula. Para testar sua função, você pode utilizar os dois poemas abaixo, cuja saída esperada é: ['and', 'are', 'is', 'you'] 

_The rose is a rose,  
And was always a rose.  
But the theory now goes  
That the apple’s a rose,  
And the pear is, and so’s  
The plum, I suppose.  
The dear only knows  
What will next prove a rose.  
You, of course, are a rose –  
But were always a rose._  

_Roses are red,  
Violets are blue,  
sugar is sweet,  
And so are you._  



__Dicas:__  
1. não se preocupe em remover a pontuação do texto.  
1. a função sorted(lista) ordena em ordem crescente (ou lexicográfica) uma lista qualquer.  
1. a função string.lower() tranforma todos os caracteres da string em letra minúscula.  
1. use `re.sub(r'[^\w]','',word.lower())` para tirar a pontuação de uma palavra: 'you.' -> 'you'

__TADs__
1. Python contém um tipo `set` que guarda elementos únicos.
1. A função `sorted` ordena dados

<!-- grade_id: cell1-a00, points: 0, grade: false, solution: true -->
```python
#Exemplo:
poem1 = """The rose is a rose,
And was always a rose.
But the theory now goes
That the apple’s a rose,
And the pear is, and so’s
The plum, I suppose.
The dear only knows
What will next prove a rose.
You, of course, are a rose –
But were always a rose."""

poem2 = """Roses are red,
Violets are blue,
sugar is sweet,
And so are you."""

import re

def common_words(text1, text2):
    text1 = text1.split() # separa as palavras do texto
    text1_split = [] # incializa a lista
    text2 = text2.split()
    text2_split = []

    for word in text1:
      # retira a pontuação das palavras e as insere numa lista
      text1_split.append(re.sub(r'[^\w]','',word.lower()))
    for word in text2:
      text2_split.append(re.sub(r'[^\w]','',word.lower()))
    
    # verifica quais palavras são iguais
    common_set = set(text1_split).intersection(text2_split)

    # ordena o set
    return (sorted(common_set))
```

<!-- grade_id: cell1t1-a00, points: 1, grade: true, solution: false -->
```python
from numpy.testing import *
"""Check that common_words returns the correct output for several inputs"""

poem1 = """The rose is a rose,
And was always a rose.
But the theory now goes
That the apple’s a rose,
And the pear is, and so’s
The plum, I suppose.
The dear only knows
What will next prove a rose.
You, of course, are a rose –
But were always a rose."""

poem2 = """Roses are red,
Violets are blue,
sugar is sweet,
And so are you."""

assert_equal(common_words(poem1, poem2), ['and', 'are', 'is', 'you'])
```

<!-- grade_id: cell1t2-a00, points: 1, grade: true, solution: false -->
```python

```

# Exercício 02:

Crie uma função que, dado uma lista de tuplas representando pares de palavras com sua respectiva traduação e uma frase, retorne a frase traduzida. As palavras que não possuem uma tradução devem ser impressas como aparecem na frase de entrada. Para testar sua função, você pode utilizar a entrada abaixo, cuja saída esperada é: o cara tossiu.

Lista de tuplas: [("galaxy", "cara tossiu"), ("kagayaku", "canalha do"), ("atsuki", "alto que"), ("yuushi", "util")]  
Frase: "o galaxy"


__Observação:__  
Esse problema, conhecido como O Fantástico Jaspion! foi adaptado de https://br.spoj.com/problems/JASPION/.

__TADs__
1. Python contém um tipo `dict/{}` que é um mapa: `chave --> valor`.
1. A função `sorted` ordena dados

<!-- grade_id: cell2-a00, points: 0, grade: false, solution: true -->
```python
#Exemplo:
lista_de_tuplas = [("galaxy", "cara tossiu"), ("kagayaku", "canalha do"), ("atsuki", "alto que"), ("yuushi", "util")]
frase = "o galaxy"


def jaspion(tuples_list, sentence):

    dictionary = dict(tuples_list) # tarnsforma a lista em dicionário
    words = sentence.split() # separa as palavras do texto

    translated = ""

    for word in words:
      if word in dictionary:
        # se a palavra está no dicionário, pega a sua tradução
        translated += dictionary.get(word) + " "
      else:
        translated += word + " "
    
    # remove o último espaço em branco
    return translated.rstrip()
```

<!-- grade_id: cell2t1-a00, points: 1, grade: true, solution: false -->
```python
"""Check that jaspion returns the correct output for several inputs"""

tl1 = [("galaxy", "cara tossiu"), ("kagayaku", "canalha do"), ("atsuki", "alto que"), ("yuushi", "util")]
s1 = "o galaxy"

assert_equal(jaspion(tl1, s1), "o cara tossiu")
```

<!-- grade_id: cell2t2-a00, points: 1, grade: true, solution: false -->
```python

```

# Exercício 03:

Crie uma função que, dado duas listas sendo uma com nomes e outras com números, retorne uma __lista__ com os nomes associados aos números ordenados em ordem decrescente. Caso mais de um nome esteja associado a um mesmo número, eles devem ser ordenados em ordem lexicográfica. Por exemplo, dado uma lista com o nome de alguns jogadores e outra lista com os salários desses jogadores, a lista com o nome dos jogadores deve ser ordenada de forma que os primeiros jogadores sejam aqueles que possuem os maiores salários. No exemplo dado, a saída esperada é: ['Messi', 'Cristiano Ronaldo', 'Neymar', 'Bale', 'Piqué', 'Kroos', 'Sergio Ramos', 'Ibrahimovic', 'Griezmann', 'Suarez', 'Iniesta', 'Lavezzi', 'Sanchez', 'Benzema', 'Muller', 'Cavani', 'Lewandowski', 'Oscar', 'Pogba', 'Neuer']


__Dicas:__  
não se esqueça de manter a relação existente entre o nome e o número, ou seja, o nome que ocupa a primeira posição da lista está associado ao número da primeira posição da outra lista. 

__Funções:__
1. sorted(dados) ordena uma lista de valores.
   Exemplo, `list(sorted([3, 2, 1]))` retorna `[1, 2, 3]`
1. `zip(lista1, lista2)`, um join de duas listas.
   Exemplo, `zip(['b', 'a'], [1, 2])` retorna `[('b', 1), ('a', 2)]`
1. `list(sorted([('b', 1), ('a', 2)]))` --> `[('a', 2), ('b', 1)]`

<!-- grade_id: cell-0c77467a22e9621e, points: 0, grade: false, solution: true -->
```python
#Exemplo:
player = ["Neymar", "Piqué", "Benzema", "Neuer", "Iniesta", "Sergio Ramos", "Messi", "Oscar", "Griezmann", "Cavani", "Bale", "Pogba", "Cristiano Ronaldo", "Sanchez", "Kroos", "Lewandowski", "Lavezzi", "Ibrahimovic", "Muller", "Suarez"]
player_salary = [81.5, 29, 23.5, 20.8, 25.5, 27.5, 126, 22.2, 26, 22.5, 44, 22, 96, 23.55, 28.2, 22.2, 25.3, 26.8, 23, 26]


def sort_lists(list1, list2):
    
    # faz o join das duas listas e as transforma em um dicionário
    players_and_salaries = dict(zip(list1, list2))

    # ordena de acordo com a ordem lexicográfica e, então, de acordo com o salário (em ordem decrescente)
    players_and_salaries = sorted(sorted(players_and_salaries), key=players_and_salaries.__getitem__, reverse=True)
    
    return players_and_salaries
```

<!-- grade_id: cell-82d1e69f4a30f889, points: 1, grade: true, solution: false -->
```python
"""Check that sort_lists returns the correct output for several inputs"""

player = ["Neymar", "Piqué", "Benzema", "Neuer", "Iniesta", "Sergio Ramos", "Messi", "Oscar", "Griezmann", "Cavani", "Bale", "Pogba", "Cristiano Ronaldo", "Sanchez", "Kroos", "Lewandowski", "Lavezzi", "Ibrahimovic", "Muller", "Suarez"]
player_salary = [81.5, 29, 23.5, 20.8, 25.5, 27.5, 126, 22.2, 26, 22.5, 44, 22, 96, 23.55, 28.2, 22.2, 25.3, 26.8, 23, 26]

assert_equal(sort_lists(player, player_salary), ['Messi', 'Cristiano Ronaldo', 'Neymar', 'Bale', 'Piqué', 'Kroos', 'Sergio Ramos', 'Ibrahimovic', 'Griezmann', 'Suarez', 'Iniesta', 'Lavezzi', 'Sanchez', 'Benzema', 'Muller', 'Cavani', 'Lewandowski', 'Oscar', 'Pogba', 'Neuer'])
```

<!-- grade_id: cell-df6952de2c0dcfc3, points: 1, grade: true, solution: false -->
```python

```

# Exercício 04:

Vamos aprender um pouco de numpy. Caso necessite, leia a documentação em http://numpy.org. Inicialmente, crie uma função que recebe duas listas de numéros, converte as duas para um vetor numpy usando `np.array` e retorna o produto interno das duas listas. 


__Dicas:__  
1. Tente fazer um código sem nenhum **for**! Ou seja, numpy permite operações em vetores e matrizes, onde: `np.array([1, 2]) + np.array([2, 2]) = np.array([3, 4])`.

__Funções:__
1. `np.sum(array)` soma os elementos do array.

<!-- grade_id: cell-81d5bf780d85f0e1, points: 0, grade: false, solution: true -->
```python
import numpy as np

#Exemplo:
x1 = [2, 4, 8]
x2 = [10, 100, 1000]


def inner(x1, x2):
    
    x1 = np.array(x1)
    x2 = np.array(x2)

    return np.dot(x1, x2)
```

<!-- grade_id: cell-fa0332fa64ee56dc, points: 1, grade: true, solution: false -->
```python
x1 = [2, 4, 8]
x2 = [10, 100, 1000]
assert_equal(20 + 400 + 8000, inner(x1, x2))
```

<!-- grade_id: cell-45249f4339f07aa8, points: 1, grade: true, solution: false -->
```python

```

# Exercício 05:

Por fim, implemente uma função utilizando numpy que recebe duas matrizes, multiplica as duas e retorne o valor médio das células da multiplicação. Por exemplo, ao multiplicar:

```
[1 2]
[3 4] 

com 

[2 1]
[1 2]

temos

[4  5 ]
[10 11]

onde a média de [4, 5, 10, 11] é

7.5, sua resposta final!
```


__Dicas:__  
1. Use o operador @ para multiplicar matrizes!

<!-- grade_id: cell-284de1851443c7cf, points: 0, grade: false, solution: true -->
```python
def medmult(X, Y):
    
    multiply = X @ Y
    return np.mean(multiply)
```

<!-- grade_id: cell-b91c763e84aeda96, points: 1, grade: true, solution: false -->
```python
X = np.array([1, 2, 3, 4]).reshape(2, 2)
Y = np.array([2, 1, 1, 2]).reshape(2, 2)
assert_equal(7.5, medmult(X, Y))
```

<!-- grade_id: cell-16fd7395e0e257d4, points: 1, grade: true, solution: false -->
```python

```

```python

```